**General instructions**:
- Submit via Jupyter Lab server
- Possible questions should be asked via the Panda course forum (last minute questions are unlikely to be answered)
- Collaborations between students during problem-solving phase on a discussion basis is OK
- However: individual code programming and submissions per student are required
- Code sharing is strictly prohibited
- We will run checks for shared code, general plagiarism and AI-generated solutions
- Any fraud attempt will lead to an auto fail of the entire course
- Do not use any additional packages except for those provided in the task templates
- Remove any plot code that was not explicitly asked for before submission. However, you are encouraged to create plots of your data for your own understanding.

# Demo exercise

This is a sample task that will serve as a test run for you for the procedure of future home exercises.
The points scored here will NOT be included in the overall statistics.
Nevertheless, you should deal with the tasks sufficiently in order not to lose any points due to formal errors in the upcoming tasks.

### Multiple Choice example - (1 point)

What is the result of 5+5 ?
- 1: 0
- 2: 10
- 3: 20

Assign the correct value to the variable *result* --> 1,2 or 3

<!-- grade_id: cell-ac32799271543259, points: 0, grade: false, solution: true -->
```julia
result = 2
```

Output:
```
2
```

<!-- grade_id: cell-27034c25518cf545, points: 1, grade: true, solution: false -->
```julia
# This is a test cell. Some tests are public, so you can see them during the exercise. They mostly contain test that check format aspects.
@assert isa(result, Number)
@assert result > 0 && result < 4

#There are also hidden tests that you cannot see, that will check wether your solution is correct
```

### Function example - (1 point)

The function test_function should return the value of the given parameter $t$ (identity function)

test_function(1) = 1, test_function(2) = 2, test_function(3) = 3, ...

<!-- grade_id: cell-5ad46beee8026867, points: 0, grade: false, solution: true -->
```julia
function test_function(t)
    return t
end
```

Output:
```
test_function (generic function with 1 method)
```

<!-- grade_id: cell-40551f84df942a1a, points: 1, grade: true, solution: false -->
```julia
# Public test
@assert isa(test_function, Function)
@assert isa(test_function(1), Number)

#Private tests
```

### Matrix example - (1 point)

The function test_function should return $t$ (parameter value) times the 2x2 unit matrix

test_matrix(1) =  $\begin{bmatrix} 1 & 0 \\ 0 & 1 \end{bmatrix}$,
test_matrix(2) =  $\begin{bmatrix} 2 & 0 \\ 0 & 2 \end{bmatrix}$,
test_matrix(3) =  $\begin{bmatrix} 3 & 0 \\ 0 & 3 \end{bmatrix}$, ...


<!-- grade_id: cell-46e04b69a5c6c3e1, points: 0, grade: false, solution: true -->
```julia
function test_matrix(t)
    return [t 0; 0 t]
end
```

Output:
```
test_matrix (generic function with 1 method)
```

<!-- grade_id: cell-91e28ae7692f64e8, points: 1, grade: true, solution: false -->
```julia
# Public test
@assert isa(test_matrix, Function)
@assert isa(test_matrix(1), AbstractMatrix)

#Private tests
```

```julia

```